# Tarea con máxima nota 15 puntos
## Estudiante: 7.5/15 puntos

<!-- grade_id: cell-1, points: 0, grade: false, solution: true -->
```python
def suma(a, b):
    # Esta función está bien
    return a + b
```

<!-- grade_id: test-1, points: 5, grade: true, solution: false -->
```python
# TEST EJERCICIO 1 - 5 PUNTOS
assert suma(2, 3) == 5
assert suma(-1, 1) == 0
assert suma(0, 0) == 0
print("OK! Ejercicio 1 correcto: 5/5 puntos")
```

<!-- grade_id: cell-2, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # Esta función tiene ERROR - range incorrecto
    if n == 0:
        return 1
    result = 1
    for i in range(1, n):  # ERROR: debería ser n+1
        result *= i
    return result
```

<!-- grade_id: test-2, points: 10, grade: true, solution: false -->
```python
# TEST EJERCICIO 2 - 10 PUNTOS
assert factorial(0) == 1
assert factorial(1) == 1
# Los siguientes tests fallarán debido al error
try:
    assert factorial(5) == 120
    print("Test 3: OK!")
except:
    print("Test 3: ERROR! - factorial(5) debería ser 120")
    
print("Ejercicio 2: 2.5/10 puntos")
print("Total estimado: 7.5/15 puntos")
```