# Tarea con máxima nota 10 puntos
## Estudiante: 8/10 puntos

<!-- grade_id: cell-1, points: 0, grade: false, solution: true -->
```python
def es_par(n):
    # Esta función está bien
    return n % 2 == 0
```

<!-- grade_id: test-1, points: 3, grade: true, solution: false -->
```python
# TEST EJERCICIO 1 - 3 PUNTOS
assert es_par(2) == True
assert es_par(3) == False
assert es_par(0) == True
print("OK! Ejercicio 1 correcto: 3/3 puntos")
```

<!-- grade_id: cell-2, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    # Esta función CASI perfecta - pequeño error
    if n <= 0:
        return 0
    elif n == 1:
        return 1
    else:
        a, b = 0, 1
        for _ in range(2, n + 1):
            a, b = b, a + b
        return a  # ERROR: debería ser b, no a
```

<!-- grade_id: test-2, points: 7, grade: true, solution: false -->
```python
# TEST EJERCICIO 2 - 7 PUNTOS
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(5) == 5  # Este pasará
# Este test fallará debido al pequeño error
try:
    assert fibonacci(6) == 8
    print("Test 4: OK!")
except:
    print("Test 4: ERROR! - fibonacci(6) debería ser 8")
    
print("Ejercicio 2: 5/7 puntos")
print("Total estimado: 8/10 puntos")
```