<!-- grade_id: squares, points: 1, grade: true, solution: true -->
```python
def squares(n):
    if n < 1:
        raise ValueError("n debe ser mayor o igual que 1")
    return [i**2 for i in range(1, n+1)]
```

<!-- grade_id: sum_of_squares, points: 1, grade: true, solution: true -->
```python
def sum_of_squares(n):
    return sum(squares(n))
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(3) == [1, 4, 9]
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("No lanzó ValueError con n=0")
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(3) == 14
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares no usa squares")
finally:
    squares = orig_squares
```

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
$$\sum_{i=1}^n i^2 = \frac{n(n+1)(2n+1)}{6}$$

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
def numero_piramidal(n):
    return sum_of_squares(n * 2)
```